## Assignment 1

#### Please copy and paste your code in the functions below.
#### Please do not add or remove any cell.

Student 1 Name: Gyuseung Sean Hwang

```python
import os
import string
```

<!-- grade_id: cell-4486cbb7b54842e9, points: 0, grade: false, solution: true -->
```python
def cleanup(intext):
    '''
    cleanup method cleans up the intext string of punctuation,
    numbers and stop words etc. and returns a lower case string.
    1. Replace all occurences of "!" with the word "ExclamationMark" (already done for you)
    2. Replace all occurences of "?" with the word "QuestionMark" (already done for you)
    3. Replace all punctuations with empty string. Hint: use string.punctuation
        Note: words like "don't" change into "dont"
    4. Replace all digits with empty string. Hint: use string.digits
    5. Make sure all words are seperated by a single space.
    6. Remove all stop words (already defiined for you)
    7. Return the cleaned up string.
    '''
    
    stopWords = {'a', 'br', 'this', 'an', 'be', 'are', 'hers', 'able', 'about', 'its', 'then', 'in', 'is', 'at', 'from', 'am', 'his', 'and', 'these', 'him', 'he', 'of', 'them', 'on', 'to', 'for', 'than', 'by', 'since', 'as', 'that', 'their', 'it', 'the', 'her', 'there', 'or'}
    
    intext = intext.replace("!", " ExclamationMark ")
    intext = intext.replace("?", " QuestionMark ")
    
    for digit in string.digits:
        intext = intext.replace(digit, "")

    for digit in string.punctuation:
        intext = intext.replace(digit, "")
    
    ret = ""
    
    for word in intext.split():
        word = word.lower()
        if word not in stopWords:
            ret += word + " "

    return ret
```

<!-- grade_id: cell-e5a2c033c0167a99, points: 10, grade: true, solution: false -->
```python
# This is a simple test for you to test your cleanup function.
assert cleanup("Great movie! GREAT!").split() == ['great', 'movie', 'exclamationmark', 'great', 'exclamationmark']
assert cleanup("I have seen 12345 movies like this.").split() == ["i", "have", "seen", "movies", "like"]
assert cleanup("This movie is bad! I don't like it....").split() == ['movie', 'bad', 'exclamationmark', 'i', 'dont', 'like']
assert cleanup("I don't know what to say! )*&$%^#$*@(^&)").split() == ['i', 'dont', 'know', 'what', 'say', 'exclamationmark']




```

<!-- grade_id: cell-07dff50164bd24ac, points: 0, grade: false, solution: true -->
```python
def directory2Vocab(dirName, outputFileName):
    '''
    create a set that contains words in all reviews
    write the words in vocab.txt (one word per line)
    order doesn't matter
    '''
    vocabSet = set()   #start with an empty set
    
    # open a new file, outputFileName, in the “w” mode.
    fileName = outputFileName
    vocabFile = open(fileName,'w')
    
    allFiles = os.listdir(dirName)
    
    for review in allFiles:
        file = open(dirName + "/" + review, 'r')
        for string in cleanup(file.read()).split():
            vocabSet.add(string)
    
    #finally write the contents of the vocab set in the vocabfile, one word per line
    vocabList = list(vocabSet)
    for word in vocabList:
        vocabFile.write(word+'\n')
    vocabFile.close()
```

<!-- grade_id: cell-d36f6025a1bfaa4e, points: 10, grade: true, solution: false -->
```python
# Test to make sure you have the vocab files created after run.
directory2Vocab('MovieReviews_subset', 'Vocab.txt')
assert os.path.isfile('Vocab.txt') == True

vocabFile = open('Vocab.txt', 'r')
vocab = vocabFile.readlines()

### Make sure there are 1127 words in the vocab file generated from reviews in MovieReviews_subset directory
assert len(vocab) == 1127

# Test you make sure you have the following words in your vocab
assert ("all\n" in vocab) == True
assert ("zero\n" in vocab) == True

```

<!-- grade_id: cell-610d012cae422602, points: 10, grade: true, solution: false -->
```python
# Test to make sure you have punctuations, stop_words, digits removed
vocabFile = open('Vocab.txt', 'r')
vocab = vocabFile.readlines()
assert ("a\n" in vocab) == False
assert (".\n" in vocab) == False
assert ("3\n" in vocab) == False

stop_words = [' a ', ' able ', ' about ', ' am ', ' an ', ' and ', ' are ', ' as ', ' at ', ' be ', ' by ', ' for ', ' from ', ' he ', ' her ', ' hers ', ' him ', ' his ', ' in ', ' is ', ' it ', ' its ', ' of ', ' on ', ' or ', ' since ', ' than ', ' that ', ' the ', ' their ', ' them ', ' then ', ' there ', ' these ',  ' this ', ' to ']
for s in stop_words:
    assert (s.strip() + '\n' in vocab) == False

```

<!-- grade_id: cell-95b6b48ca2eca489, points: 0, grade: false, solution: true -->
```python
def directory2Features(dirName, vocabfilename):
    '''
    create feature vector for each review
    Format each feature vector as following:
    rating word1_index:word1_count word2_index:word2_count word3_index:word3_count ...
    index must be in increasing order
    
    for example:
    4 0:3 3:2 10:2 33:1 53:2 78:1 241:2
    
    write all feature vectors to file <dirName>_fVectors.txt
    '''
    
    # open and read the vocabfilename
    # put each word in vocabfile as an element in a list.
    vocFile = open(vocabfilename, 'r') #Open the vocab file
    vocabWords = vocFile.read().splitlines()
    
    # open a new file, <dirName>_fVectors.txt, in the “w” mode.
    vecFileName = dirName.split("/")[0]+'_fVectors.txt'
    fVectors = open(vecFileName, 'w')
    
    allFiles = os.listdir(dirName)
    for f in allFiles:
        #open file in read format
        #read file as a string
        file = os.path.join(dirName,f)
        reviewFile = open(file, 'r')
        allContent = reviewFile.read()
        
        # 1. write code to clean up the string by calling cleanup method
        allContent = cleanup(allContent)
        # 2. extract rating from the file name, f.
        rating = f.split("_")[0]
        # 3. Create the feature vector to the rating
        fVector = rating
        counter = {}
        for word in allContent.split():
            wordIndex = vocabWords.index(word)
            if wordIndex in counter:
                counter[wordIndex]+=1
            else:
                counter[wordIndex]=1
            
        for wIndex in sorted(counter.keys()):
            fVector += " " + str(wIndex) + ":" + str(counter[wIndex])
        
        # 5. Append the feature vector to the rating and write it to the fVectors file
        fVectors.write(fVector + "\n")
        
    vocFile.close()
    fVectors.close()
```

<!-- grade_id: cell-54788abc11cd108e, points: 20, grade: true, solution: false -->
```python
# Test to make sure you have the fVectors files created after run.
directory2Features('MovieReviews_subset', "Vocab.txt")
assert os.path.isfile('MovieReviews_subset_fVectors.txt') == True

vecFile = open('MovieReviews_subset_fVectors.txt', 'r')
vec = vecFile.readlines()

### Make sure there are 22 lines in the fVectors file generated from reviews in MovieReviews_subset directory
### (one line per review)
assert len(vec) == 22

```

<!-- grade_id: cell-a12d9ac00597ff5e, points: 10, grade: true, solution: false -->
```python
# Test on the complete dataset
directory2Vocab('MovieReviews_full/', 'Vocab_full.txt')
assert os.path.isfile('Vocab_full.txt') == True

# Test to make sure you have punctuations, stop_words, digits removed
vocabFile = open('Vocab_full.txt', 'r')
vocab = vocabFile.readlines()
assert ("a\n" in vocab) == False
assert (".\n" in vocab) == False
assert ("3\n" in vocab) == False

```

```python
def createFeatureVectors(dirName, vocabFile):
    directory2Vocab(dirName, vocabFile)
    directory2Features(dirName, vocabFile)


```

```python
## Running createFeatureVectors("MovieReviews_Example/", "Vocab_Example.txt") generates 2 files:
## 1. Vocab_Example.txt
## 2. MovieReviews_Example_fVectors.txt

### The Vocab_Example.txt file contains 434 words
### MovieReviews_Example_fVectors.txt contains 5 lines: 1 line per review.

### These 2 files are already provided to you to confirm your results.
```

```python
createFeatureVectors("MovieReviews_Example/", "Vocab_Example.txt")
```

```python

```