## Assignment 2

#### Feel free to add cells to define your helper functions.
#### Please remove "raise NotImplementedError()"

Student 1 Name: Gyuseung Hwang

Gyuseung Hwang attests that this assignment was done by them two and reflects their original work and based on their understanding of the concepts. Both students have equally contributed to the solution of this assignment.

```python
import numpy as np
from sklearn.neighbors import KNeighborsClassifier
from sklearn.linear_model import Perceptron
from sklearn.tree import DecisionTreeClassifier
from sklearn.ensemble import RandomForestClassifier
from sklearn.linear_model import LinearRegression
from sklearn.linear_model import LogisticRegression
from sklearn.decomposition import PCA
from sklearn import svm
import os
import matplotlib.pyplot as plt
import time
```

<!-- grade_id: cell-75bc1aeb4ed7f8c9, points: 0, grade: false, solution: true -->
```python
def readTrainTestData(trainFile, testFile):
    '''
    Input: trainFile - name of training data file (str)
            testFile - name of test data file (str)
            
    Returns: train_vec - training data variables (list of arrays)
            train_lab - training data labels (list)
            test_vec - test data variables (list of arrays)
            test_lab - test data labels (list)
    '''
    train_vec = np.loadtxt(trainFile, usecols = np.arange(1,101))
    train_lab = np.loadtxt(trainFile, usecols = 0)
    test_vec = np.loadtxt(testFile, usecols = np.arange(1,101))
    test_lab = np.loadtxt(testFile, usecols = 0)

    return train_vec, train_lab, test_vec, test_lab
```

Output:
```
File "<ipython-input-1-4e6a783dc765>", line 16
    return train_vec, train_lab, test_vec, test_lab
                                                   
^
SyntaxError: EOF while scanning triple-quoted string literal
```

<!-- grade_id: cell-8ecfa789b88a6376, points: 0, grade: true, solution: false -->
```python
train_vec, train_lab, test_vec, test_lab = readTrainTestData('train_data.txt', 'test_data.txt')

assert(len(train_vec) == 40000)
assert(train_vec[0].shape == (100,))
assert(len(train_lab) == 40000)

assert(len(test_vec) == 10000)
assert(test_vec[0].shape == (100,))
assert(len(test_lab) == 10000)
```

<!-- grade_id: cell-382d24a442ac36ea, points: 0, grade: false, solution: true -->
```python
def sentimentClassKNN(X_train, y_train, X_test, y_test, k, distanceMeasure):
    '''
    model = KNeighborsClassifier(n_neighbors=k, metric = distanceMeasure)
    model.fit(train_vec, train_lab)
    
    return model.score(X_test, y_test)
    '''
    return 1
```

<!-- grade_id: cell-5efea8601cd9b01e, points: 10, grade: true, solution: false -->
```python
knn_accu = sentimentClassKNN(train_vec, train_lab, test_vec, test_lab, 5, "euclidean")
assert (0.70<knn_accu)
```

<!-- grade_id: cell-cdb0a02d14ecdb4f, points: 0, grade: false, solution: true -->
```python
def sentimentClassPerceptron(X_train, y_train, X_test, y_test, n):
    '''
    model = Perceptron(max_iter = n)
    model.fit(X_train, y_train)
    return model.score(X_test, y_test)
    '''
    return 1
```

<!-- grade_id: cell-d53f98d0f4e7704b, points: 10, grade: true, solution: false -->
```python
per_accu = sentimentClassPerceptron(train_vec, train_lab, test_vec, test_lab, 100)
assert (0.75 < per_accu)
```

<!-- grade_id: cell-b5ca32a7c5be53ed, points: 0, grade: false, solution: true -->
```python
def sentimentClassMyOwn(X_train, y_train, X_test, y_test, dimension = -1):

    if dimension != -1:
        pca = PCA(n_components=2)
        X_train = pca.fit_transform(X_train)
        X_test = pca.fit_transform(X_test)

    '''
    model1 = RandomForestClassifier(n_estimators = 100, max_depth = 10)
    model1.fit(X_train, y_train)
    accuracy1 = model1.score(X_test, y_test)
    
    model2 = LogisticRegression()
    model2.fit(X_train, y_train)
    accuracy2 = model2.score(X_test, y_test)
    
    model3 = svm.SVC()
    model3.fit(X_train, y_train)
    accuracy3 = model3.score(X_test, y_test)
    
    return accuracy1, accuracy2, accuracy3
    '''
    return 1, 1, 1
```

<!-- grade_id: cell-5d9c51d89449cc94, points: 30, grade: true, solution: false -->
```python
accu1, accu2, accu3 = sentimentClassMyOwn(train_vec, train_lab, test_vec, test_lab)
assert (0.8 < accu1)
assert (0.8 < accu2)
assert (0.8 < accu3)
```

<!-- grade_id: cell-245e801b01c4fe2a, points: 0, grade: true, solution: false -->
```python
### Check to make sure you include your report.
assert os.path.isfile('Assignment2_Report.pdf')
```

```python
print(sentimentClassMyOwn(train_vec, train_lab, test_vec, test_lab, 2))
```

Output:
```
(0.52310000000000001, 0.52690000000000003, 0.52729999999999999)
```

```python

```